## EC Notebook for Lecture 10:  Functions in Python

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30am on **Friday, Sept. 20** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.


## 1. Importing pandas and reading data
For this extra credit notebook, we will still use the Illinois GPA Distribution dataset -- as we did in lecture -- to do some exercise to get familiar with the function we learnt in the lecture

You feel free to access the "GPA" dataset here: https://github.com/wadefagen/datasets/tree/master/gpa
At first, you need to import the library pandas and reading the dataset.

<!-- grade_id: cell-b94d835c3afb8ee1, points: 0, grade: false, solution: true -->
```python
import pandas as pd
df = pd.read_csv("gpa.csv")
df

```

Output:
```
Year    Term YearTerm Subject  Number                    Course Title  \
0      2018    Fall  2018-fa     AAS     100    Intro Asian American Studies   
1      2018    Fall  2018-fa     AAS     100    Intro Asian American Studies   
2      2018    Fall  2018-fa     AAS     100    Intro Asian American Studies   
3      2018    Fall  2018-fa     AAS     100    Intro Asian American Studies   
4      2018    Fall  2018-fa     AAS     100    Intro Asian American Studies   
5      2018    Fall  2018-fa     AAS     105        Intro to Arab Am Studies   
6      2018    Fall  2018-fa     AAS     211    Asian Americans and the Arts   
7      2018    Fall  2018-fa     AAS     246    Asian American Youth in Film   
8      2018    Fall  2018-fa     AAS     287        Food and Asian Americans   
9      2018    Fall  2018-fa     ABE     100  Intro Agric & Biological Engrg   
10     2018    Fall  2018-fa     ABE     223    ABE Principles: Machine Syst   
11     2018    Fall  2018-fa     ABE     224   
```

## 2. Review: Group Function
The Student Union at Illinois would like to investiage the grading policy for the university courses and subjects， especailly for "ACCY", "CS", "MATH", "STAT".

### Question 1: 
Using the GPA dataset, you need to calculate the proportion of ("A+","A") in ("A+","A","A-","B+","B","B-") for subjects "ACCY", "CS", "MATH", "STAT", respectively. Hint: use the groupby function.

<!-- grade_id: cell-b1e450db85c95e49, points: 0, grade: false, solution: true -->
```python
df_subj = df[(df["Subject"] == "CS") | (df["Subject"] == "ACCY") | (df["Subject"] == "MATH") | (df["Subject"] == "STAT")]
group = df_subj.groupby("Subject")
df_group = group.agg("sum").reset_index()
df_group
ratio = (df_group["A+"] + df_group["A"]) / (df_group["A+"] + df_group["A"] + df_group["B+"] + df_group["B-"] + df_group["B"] + df_group["A-"])
print(ratio[1])

```

Output:
```
0.4851959027145017
```

<!-- grade_id: cell-fc6499a6e7186f28, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(ratio[1] == 0.4851959027145017), "The ratio for the ACCY does not appear correct."
```

## 3. Review: Group Function
Besides the raito of ("A+","A") in ("A+","A","A-","B+","B","B-"), the student union also wants to know the difference between the largest number of "A+" and the largest number of "A" in each subject, and what is the largest gap among all subjects. 

### Question 2: 
Using the GPA dataset, you need to find the difference between the largest number of "A+" and the largest number of"A" for each subject firstly, and then based on what you have, to calculate the largest gap among all subjects. 

<!-- grade_id: cell-17f1a17a9291bd92, points: 0, grade: false, solution: true -->
```python

df_group_max = group.agg("max").reset_index()

gap = df_group_max["A+"] - df_group_max["A"]
print(gap_largest)

```

Output:
```
50
```

<!-- grade_id: cell-63f5df829092ffa4, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(gap_largest == 481), "The number does not appear correct."
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-24-cb0a68c4f845> in <module>
      4 # - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
      5 
----> 6 assert(gap_largest == 481), "The number does not appear correct."

AssertionError: The number does not appear correct.
```

## 4. Review: Group Function
In addition, the student union also wants to know hom many subjects which the instructors give "A" but not "A+". 


### Question 3: 
Using the GPA dataset, we have defined a function to give the label (column "course_type") for those course which the instructors give "A" but not "A+". For your work, you need to decide how many subjects which the instructors give "A" but not "A+". Hint: The question is asking about each subject not each course, so you also need to use groupby function.

<!-- grade_id: cell-656278288cff2886, points: 0, grade: false, solution: true -->
```python
## Just Run the following code. 
def course(x):

   if x["A+"] == 0 and x["A"] > 0 : return 1

   else: return 0
    
df["course_type"] = df.apply(course, axis=1)
## Just Run the above code. 

num = ...
print(num)

```

<!-- grade_id: cell-ce3a059c412e2e95, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Question 1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(num == 13), "The number of this type of course does not appear correct."
```