# EC Notebook for Lecture 6: Working with `len()`

This extra credit Python notebook will let you practice the material you saw in lecture.  Completing all parts of this notebook will earn +1 extra credit point to your grade in STAT 107! :)

This notebook is worth +1 if turned in before 11:30am on **Wednesday, Sept. 11** *(30 minutes before the next STAT 107 lecture)*.  You can feel free to complete it anytime for extra practice.

## 1. Loading pandas and importing data

For this notebook, we will use the "Course Catalog" dataset that is our dataset of the week.

Here's the link: https://docs.google.com/spreadsheets/d/1ZnpZpipyui_zIIBblmalRYiyEG1glUKinqJtQxVL04o/edit#gid=0

Use the following cell to load the pandas library and read the DataSet as a DataFrame:

<!-- grade_id: cell-0c0abea61271817c, points: 0, grade: false, solution: true -->
```python

```

## 2. Courses in STAT, CS, and IS

This course, Data Science Discovery, is cross-listed as **STAT 107**, **CS 107**, and **IS 107** as this is a truly cross-disciplinary course.

Using the "Course Catalog" dataset, find the number of courses is the subjects **STAT**, **CS**, and **IS** and store the value in the cells `STAT_courses_ct`, `CS_courses_ct`, and `IS_courses_ct`,

- *Make sure that these values are **numbers**, not DataFrames!*

<!-- grade_id: cell-18acb525a86ffa3e, points: 0, grade: false, solution: true -->
```python


STAT_courses_ct = ...
print(f"There are a total of {STAT_courses_ct} unique STAT courses in the course catalog.")


```

<!-- grade_id: cell-304c0e604514f773, points: 0, grade: false, solution: true -->
```python


CS_courses_ct = ...
print(f"There are a total of {CS_courses_ct} unique CS courses in the course catalog.")


```

<!-- grade_id: cell-80924a83e9d8942e, points: 0, grade: false, solution: true -->
```python


IS_courses_ct = ...
print(f"There are a total of {IS_courses_ct} unique IS courses in the course catalog.")


```

<!-- grade_id: cell-8f911f7a0d33fd01, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( STAT_courses_ct == 57 ), "The number of STAT courses appears to be incorrect."
assert( CS_courses_ct == 118 ), "The number of STAT courses appears to be incorrect."
assert( IS_courses_ct == 92 ), "The number of STAT courses appears to be incorrect."
```

## Submit Your Work!

Make sure to **Save and Checkpoint** your notebook, exit Jupyter, and submit your work! :)